<problem display_name="Sandbox World" markdown="null">
  <script type="loncapa/python">

import json
def vglcfn(e, ans):
  
  par = json.loads(ans)
  
  answer = json.loads(par["answer"]);
  fat = answer["fat"];
  calories = answer["calories"];
  protein = answer["protein"];
  carbs = answer["carbs"];

  totalGrams = carbs + fat + protein;
  carbs_p = carbs / totalGrams;
  fat_p = fat / totalGrams;
  protein_p = protein / totalGrams;
  
  return (calories &lt; 800 and calories &gt; 500 and fat_p &gt; .05 and fat_p &lt; .25 and carbs_p &gt; .3 and carbs_p &lt; .5 and protein_p &gt; .35 and protein_p &lt; .55)


 

</script>
  <script type="text/javascript">

$('iframe').css('height','500');
$('iframe').css('width','700');
$('iframe').css('border','10px solid white');
$('iframe').css('box-shadow','0px 0px 80px gray');


var opened = false;
var _to = null;

var otop = $('iframe').offset().top;
var oleft = $('iframe').offset().left;
var owidth = $('iframe').width();
var oheight = $('iframe').height()

$('iframe').mouseover(function()
{

   if(_to)
          {
            window.clearTimeout(_to);
            _to=null;
          }

      _to = window.setTimeout(function(){


          $('iframe').css('position','fixed');
          $('iframe').css('z-index','1000000');
          $('iframe').css('box-shadow','0px 0px 180px gray');
          $('iframe').css('top',otop)
          $('iframe').css('left',oleft)
          $('iframe').css('width',owidth)
          $('iframe').css('height',oheight)

          $('iframe').animate({
            'top':'15%',
            'left':'15%',
            'width':'70%',
            'height':'70%'});

         

        },1500);
        
      
});
$('iframe').mouseout(function()
{
          if(_to)
          {
            window.clearTimeout(_to);
            _to=null;
          }
          _to = window.setTimeout(function(){
            $('iframe').animate({
            'top':otop,
            'left':oleft,
            'width':owidth,
            'height':oheight},function(){
          $('iframe').css('top','')
          $('iframe').css('left','')
          $('iframe').css('height','500');
          $('iframe').css('width','700');
          $('iframe').css('position','');
          $('iframe').css('z-index','');
          $('iframe').css('box-shadow','0px 0px 80px gray');
              });
        
        },1500);
         
          
        
});

</script>
  <p>
Complete the Nutrition Mission Game.
</p>
  <customresponse cfn="vglcfn">
    <jsinput gradefn="WebGLDemo.getGrade" get_statefn="WebGLDemo.getState" set_statefn="WebGLDemo.setState" initial_state="{}" width="800" height="600" html_file="http://ec2-54-173-64-18.compute-1.amazonaws.com:3000/adl/sandbox/akJpH8kSqyaXwLiB/" sop="false"/>
  </customresponse>
</problem>
